# Python Addition Test
Simple nbgrader assignment to test code cell extraction

## Exercise 
add the two variables `x` and `y` and save their sum in the variable `z`

<!-- grade_id: cell-9fb6fe3588ec4909, points: 0, grade: false, solution: true -->
```python
x = 3
y = 7
### BEGIN SOLUTION
z = x + y
### END SOLUTION
```

<!-- grade_id: cell-744e5dbe470759ae, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert z == 10
### END HIDDEN TESTS
```

```python

```